# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

### 21.1 What does the keyword `def` mean before a function name?

> A) _define_ - You are saying that this is the function definition
>
> B) _default_ - This is the default work the function should do
>
> C) _def_ - Doesn't stand for anything, it's just what you put there
>
> D) _defer_ - Meaning don't do anything right now

```python
answer211 = "A"
```

### 21.2 What is the minimum number of parameters that you have to provide for this function to work?

```python
def my_function (fname='John', lname='Doe', gender='M'):
    ...
```

> A) 3
>
> B) 2
>
> C) 1
>
> D) 0

```python
answer212 = "A"
```

### 21.3 Which of the following lines of code does not have any functions in it?


> A) `value = math.log(23,2)`
>
> B) `pi * (radius ** 2)`
>
> C) `[ x for x in range(10) ]`
>
> D) `print('I am a good student')`

```python
answer213 = "D"
```

### 22.1 Which is true about docstrings?

> A) They are stored in Python as just a long string inside the function definition.
>
> B) They can contain tests to help you demonstrate your code works correctly.
>
> C) They tell Python what to say when someone asks for `help()` on your function.
>
> D) All of these

```python
answer222 = "D"
```

### 22.2 True or False: From now on you should include docstrings in the _applied_ homework assignments?

Type `answer222 = True` or `answer222 = False` as your response

```python
answer222 = True
```

### 23.1 Which of the following are typical ways of encapsulating logical rules or mapping in Python?

> A) dict
>
> B) str
>
> C) function
>
> D) A and C
>
> E) None of these

```python
answer231 = "D"
```

### 23.2 What does DRY stand for?


> A) Do not repeat yourself
>
> B) Do not repeat yourself
>
> C) Do not repeat yourself
>
> D) All of the above

```python
answer232 = "D"
```

### 24.1 When Python tells you there's been an exception, what is the key word that we use in our Python code to indicate how Python should react to that exception?

> A) Throw
>
> B) Raise
>
> C) Except
>
> D) Crash

```python
answer241 = "B"
```

### 24.2 When you, as the programmer, want to notify Python that an exception should be triggered, what is the key word we use for that?

> A) Throw
>
> B) Raise
>
> C) Except
>
> D) Crash

```python
answer242 = "B"
```

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(strings):
    """(list) -> string
    we run a loop to keep a memory of longest string length and return position of the 
    string which has max length in list.
    If the same element is present more than once, the method returns the index of the 
    first occurrence of the element.
    
    >>> strings = ['happy', 'birthday', 'to', 'me']
    >>> longest(strings)
    1
    """
    max_len = -1
    for num in strings: 
        if len(num) > max_len: 
            max_len = len(num) 
            result = num
            pos = strings.index(result)
    return pos

### END SOLUTION
```

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
longest("happy birthday to me".split(" "))
```

Output:
```
1
```

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
longest("enjoy class".split(" "))
```

Output:
```
0
```

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
longest(['when','what','where','how','who'])
```

Output:
```
2
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

### 25.2 Exception Handling

Write a function called `safe_average(L)` to calculate the average of a list of values.  If any of the values in the list are not valid numbers, then the function should simply ignore them and exclude from the average.  If none of the values are valid numbers, then the function should return None.

Example:

```
>>> safe_average([100,200,'three'])
150

>>> safe_average(['one','two','three'])
None
```

<!-- grade_id: exception-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

from numbers import Number

def safe_average(n):
    """(list) -> float
    This function will calculate the average of the list of all input value. 
    The input is expected to be a valid number, if they
    are not valid numbers, then simply ignore them and exclude from the average.
    If none of the values are valid numbers, then the function should return None.
    
    >>> safe_average([150,140,240,'three',340])
    217.5

    >>> safe_average(['one', 'five', 'ten'])
    None
    
    """
    
    if all(isinstance(item, str) for item in n):
        return None
    else:
        temp = list(filter(lambda x: isinstance(x, Number), n))

    result = 0
    for i in temp:
        result += i
        ave_num = result / len(temp)

    return ave_num  

### END SOLUTION
```

<!-- grade_id: exception-answer, points: 0, grade: false, solution: true -->
```python
safe_average([1,2,3])
```

Output:
```
2.0
```

<!-- grade_id: exception-answer, points: 0, grade: false, solution: true -->
```python
safe_average([1,'two',3])
```

Output:
```
2.0
```

<!-- grade_id: exception-answer, points: 0, grade: false, solution: true -->
```python
safe_average(['one','two','three'])
```

<!-- grade_id: exception-tests, points: 10, grade: true, solution: false -->
```python
assert safe_average([1,2,3]) == 2
assert safe_average([1,'two',3]) == 2
assert safe_average(['one','two','three']) == None
```

---

## Rubric for 25.1 and 25.2

* Create a properly working function
* Follow good practices for naming the function and its parameters
* Provide meaningful documentation
* Provide sufficient docstring tests