# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
# Using a magic function
%matplotlib inline
# importing pandas as pd
import pandas as pd

# Reading the data by the read_csv() method in pandas
hhcaps = pd.read_csv('/data/hhcaps.csv')
```

## #01 - 

How many columns does this file contain?

```python
# Printing out the shape(rows, columns) of the dataframe
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
# The hhcaps.columns returns an array of all column names then we are just computing it's length by len() function
answer = len(hhcaps.columns)

```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

```python

```

```python
type(answer)
```

Output:
```
int
```

```python
answer
```

Output:
```
39
```

```python

```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
# unique() helps us to get all unique elements
# len() return the length 
answer = hhcaps["State"].unique()
answer = len(answer)

```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python

```

```python
type(answer)
```

Output:
```
int
```

```python
answer
```

Output:
```
55
```

```python

```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
# Grouping the dataframe by "State" column
data = hhcaps.groupby( by="State" )
# Creating two empty lists on purpose
occ = []
info = []
# looping through the data and appending the needed data into the above lists
for state, df in data:
    info.append(df.shape[0])
    occ.append((state, df.shape[0]))
# Getting out the maximum value from the list called "info"
maximum = max(info)
# Again looping though the "occ" list to get the certain answer
for i, j in occ:
    if j == maximum:
        val = i

# Assigning the certain answer called "val" to the answer
answer = val
```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

```python

```

```python
type(answer)
```

Output:
```
str
```

```python
answer
```

Output:
```
'TX'
```

```python

```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
# Again grouping the dataframe by "State" column
data = hhcaps.groupby(by = "State")
# Creating two empty lists on purpose
score = []
vals = []
# Appending certain data to the empty lists
for i, j in data:
    vals.append(j["Star Rating for health team communicated well with them"].mean())
    score.append((i, j["Star Rating for health team communicated well with them"].mean()))
# the max() returns back the the maximum value from the list called "vals"
maximum = max(vals)
for i, j in score:
    if j == maximum:
        state = i
        
# Assigning the certain answer called "state" to the answer
answer = state

```

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

```python

```

```python
type(answer)
```

Output:
```
str
```

```python
answer
```

Output:
```
'ME'
```

```python

```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
# getting the data from the dataframe which matches certain condition
data = hhcaps[hhcaps['Type of Ownership'] == "Hospital Based Program"]
# The mean() returns the mean value..
val = data["Star Rating for how patients rated overall care from agency"].mean()
# Assigning the answer "val" to the answer
answer = val

```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))

```

```python

```

```python
type(answer)
```

Output:
```
numpy.float64
```

```python
round(answer,5)
```

Output:
```
3.69252
```

```python

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
for column in hhcaps.columns:
    # for each column in dataframe, get number of unique values
    num_of_uniques = len(hhcaps[column].unique())
    # if number of unique values is one for any column, this is column name we need
    if num_of_uniques == 1:
        answer = column
        break
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())

```

```python

```

```python
type(answer)
```

Output:
```
str
```

```python
answer.lower()
```

Output:
```
'offers nursing care services'
```

```python

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
# Here we are using "startswith('Offers')" to match if the sentense starts with "Offers"
data = hhcaps.loc[:, hhcaps.columns.str.startswith('Offers')]
# Converting all the values to the numpy array by using "to-numpy()" method
values = data.to_numpy()
total = 0
for i in range(0, len(values)):
    if all(values[i]) == True:
        total += 1
answer = total

```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

```python

```

```python
type(answer)
```

Output:
```
int
```

```python
answer
```

Output:
```
9029
```

```python

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
answer = pd.to_numeric(hhcaps["Number of completed Surveys"], errors="coerce").median()

```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)

```

```python

```

```python
type(answer)
```

Output:
```
numpy.float64
```

```python
answer
```

Output:
```
84.0
```

```python

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
# get numbr of row with have "SAINT LOUIS" as value in `City` column
answer = len([x   for x in hhcaps["City"]  if "SAINT LOUIS" == x.strip()])
answer
```

Output:
```
23
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

```python

```

```python
type(answer)
```

Output:
```
int
```

```python
answer
```

Output:
```
23
```

```python

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'
# get all rows whose `Percent of patients ...` column is numeric and create new dataframe
# filter the resultant dataframe, such that it has only `State` and `Percent of patients ...` columns
df = hhcaps[hhcaps[field].str.isnumeric()].filter(["State", field], axis=1)

# convert `Percent of patients ...` column to integer
df[field] = df[field].astype(int)

# group `State` column by finding mean of `Percent of patients ...` column and reset index.
df = df.groupby(by="State")[field].mean().reset_index()

# get the `State` with the maximum value for `Percent of patients ...` column
answer = df[df[field] == df[field].max()]["State"].tolist()[0]

```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```

```python
type(answer)
```

Output:
```
str
```

```python
answer
```

Output:
```
'MS'
```

```python

```

```python

```