# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
›# importing required libraries
import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
# First, put in your user name below
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!

MYTABLE='srai' + '_data'


# Then, create a connection to the same database, using the 
# same user name and password in the `mysql-examples.ipynb` 
# file.  Use conn as the database connection variable.

host = 'slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2020'
password = 'SLUcor2020'
database = 'hds5210'

conn = create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + '/' + 
                     database, echo=False)
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'paulboal_data')
assert(conn.connect())
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
MYTABLE
```

Output:
```
'srai_data'
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
conn.connect()
```

Output:
```
<sqlalchemy.engine.base.Connection at 0x7fa70532cda0>
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python

```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'

### SOLUTION GOES HERE

data = pd.read_csv(url)

# CREATE TABLE and fill it with the contents of the dataframe `data`
# the table scheme will be created automatically by pandas
data.to_sql(con=conn, name=MYTABLE, if_exists='replace')
```

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
```

Output:
```
index                                 Provider_Name         NPI    CCN  \
0        0                          Sutter Bay Hospitals  1659439834  50008   
1        1  PRIME HEALTHCARE SERVICES - GARDEN GROVE LLC  1659538858  50230   
2        2                        ST MARY MEDICAL CENTER  1669456299  50300   
3        3                     MADERA COMMUNITY HOSPITAL  1669673646  50568   
4        4                  Temecula Valley Hospital Inc  1679816201  50775   
..     ...                                           ...         ...    ...   
323    323                   KAISER FOUNDATION HOSPITALS  1326119967  50071   
324    324      COUNTY OF LOS ANGELES AUDITOR CONTROLLER  1336154020  50717   
325    325                    REEDLEY COMMUNITY HOSPITAL  1336167550  50192   
326    326  PROVIDENCE HEALTH SYSTEM-SOUTHERN CALIFORNIA  1336173269  50235   
327    327                   KAISER FOUNDATION HOSPITALS  1336294040  50411   

    Medicaid_EP_Hospital_Type           Street_Address  
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
xyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (328,20))
assert(list(dxyz.columns) == ['index', 'Provider_Name', 'NPI', 'CCN',
       'Medicaid_EP_Hospital_Type', 'Street_Address', 'City', 'County',
       'State', 'Zip_Code', 'Payment_Year_Number', 'Program_Type',
       'Total_payments', 'Last_Program_Year', 'Last_Payment_Year',
       'Last_Payment_Criteria', 'Most_Recent_Disbursement_Amount', 'LONGITUDE',
       'LATITUDE', 'FID'])
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz.shape
```

Output:
```
(328, 20)
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
list(dxyz.columns)
```

Output:
```
['index',
 'Provider_Name',
 'NPI',
 'CCN',
 'Medicaid_EP_Hospital_Type',
 'Street_Address',
 'City',
 'County',
 'State',
 'Zip_Code',
 'Payment_Year_Number',
 'Program_Type',
 'Total_payments',
 'Last_Program_Year',
 'Last_Payment_Year',
 'Last_Payment_Criteria',
 'Most_Recent_Disbursement_Amount',
 'LONGITUDE',
 'LATITUDE',
 'FID']
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python

```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Zip_Code` field and `population`'s `zipcode` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
answer = []


population = pd.read_sql_query("SELECT * FROM population", conn)


filtered_data = data[data["Zip_Code"].isin(population["zipcode"])]
population["zipcode"] = population["zipcode"].astype("int64")
merged_data = pd.merge(filtered_data, population, left_on="Zip_Code", right_on="zipcode")



largest_population = merged_data["population"].max()
temp_df = merged_data[merged_data["population"] == largest_population]
zipcode = list(temp_df["Zip_Code"])[0]
answer = list(map(str, temp_df[ merged_data["Zip_Code"] == zipcode ]["NPI"]))
```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
answer = []

# read the population table into dataframe
population = pd.read_sql_query("SELECT * FROM population", conn)

population
```

Output:
```
zipcode  population  median_age  total_males  total_females  \
0     91371           1        73.5            0              1   
1     90001       57110        26.6        28468          28642   
2     90002       51223        25.5        24876          26347   
3     90003       66266        26.3        32631          33635   
4     90004       62180        34.8        31302          30878   
..      ...         ...         ...          ...            ...   
314   93552       38158        28.4        18711          19447   
315   93553        2138        43.3         1121           1017   
316   93560       18910        32.4         9491           9419   
317   93563         388        44.5          263            125   
318   93591        7285        30.9         3653           3632   

     total_households  avg_household  
0                   1           1.00  
1               12971           4.40  
2               11731           4.36  
3               15642           4.22  
4   
```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
# We want the rows in `data` whose `Zip_Code` exists in `population` dataframe,
# we filter it accordingly and store in `filtered_data`

filtered_data = data[data["Zip_Code"].isin(population["zipcode"])]

population["zipcode"] = population["zipcode"].astype("int64")

# merge the `filtered_data` dataframe with `population`.
# here we are merging depending on zipcode
# result stored in `merged_data`
merged_data = pd.merge(filtered_data, population, left_on="Zip_Code", right_on="zipcode")


merged_data
```

Output:
```
Provider_Name         NPI    CCN  \
0                    LONG BEACH MEMORIAL MEDICAL  1720371065  50727   
1                AHMC GARFIELD MEDICAL CENTER LP  1730171265  50737   
2                 AHMC MONTEREY PARK HOSPITAL LP  1780676221  50736   
3      PRIME HEALTHCARE SERVICES - SAN DIMAS LLC  1740447945  50588   
4           City of Hope National Medical Center  1750358297  50146   
..                                           ...         ...    ...   
82  DOWNEY REGIONAL MEDICAL CENTER-HOSPITAL, INC  1265433551  50393   
83    AHMC SAN GABRIEL VALLEY MEDICAL CENTER, LP  1275720377  50132   
84                   KAISER FOUNDATION HOSPITALS  1295880912  50677   
85  PROVIDENCE HEALTH SYSTEM-SOUTHERN CALIFORNIA  1336173269  50235   
86                   KAISER FOUNDATION HOSPITALS  1336294040  50411   

   Medicaid_EP_Hospital_Type        Street_Address            City  \
0       Acute Care Hospitals      1720 TERMINO AVE      LONG BEACH   
1       Acute Care Hospitals    525 N GARF
```

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
# get the largest population from column

largest_population = merged_data["population"].max()

# fetch the row with the `largest_population`
temp_df = merged_data[merged_data["population"] == largest_population]

# get the zipcode from the row
zipcode = list(temp_df["Zip_Code"])[0]

# get all providers' `NPI`s who have Zip codes equal to `zipcode`
answer = list(map(str, temp_df[ merged_data["Zip_Code"] == zipcode ]["NPI"]))
```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/ipykernel_launcher.py:12: UserWarning: Boolean Series key will be reindexed to match DataFrame index.
  if sys.path[0] == '':
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
type(answer)
```

Output:
```
list
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
answer
```

Output:
```
['1194016923']
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python

```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
answer = None

# create a new dataframe with only columns we need
# `Zip_Code`
# `Medicaid_EP_Hospital_Type`
# `Total_payments`
data_copy = data[["Zip_Code", "Medicaid_EP_Hospital_Type", "Total_payments"]].copy()

# Group by `Zip_Code`, `Medicaid_EP_Hospital_Type`
# and calculate sum of the rest of the numeric columns.
# Here we only have one more numeric column, that is, `Total_payments`
temp_df = data_copy.groupby(by=["Zip_Code", "Medicaid_EP_Hospital_Type"]).sum().reset_index()

# merge the `temp_df` dataframe with `population`.
# here we are merging depending on zipcode
# result stored in `temp_df`
temp_df = pd.merge(temp_df, population, left_on="Zip_Code", right_on="zipcode")

# Calculate `Total_Payments_per_person`
# `Total_Payments_per_person` = `Total_payments` / `population `
temp_df["Total_Payments_per_person"] = temp_df["Total_payments"] / temp_df["population"]

# calulate average of `Total_Payments_per_person`
# for each `Medicaid_EP_Hospital_Type`
temp_df = temp_df.groupby(by=["Medicaid_EP_Hospital_Type"])["Total_Payments_per_person"].mean()

# Reset Index, so that we have all the columns,
# this is to retain the lost columns during groupby
answer = temp_df.reset_index()

# Now that our `Total_Payments_per_person`, contains average
# for each `Medicaid_EP_Hospital_Type`, we can rename it to
# 'Avg_Pay_per_Capita'. Therefore, the resulting column names in the dataframe will be
# `Medicaid_EP_Hospital_Type` and `Avg_Pay_per_Capita`.
answer.columns = ['Medicaid_EP_Hospital_Type', 'Avg_Pay_per_Capita']

answer
```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita
0      Acute Care Hospitals        17974.362790
1      Children's Hospitals          111.559021
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 17974.363)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 111.559)
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
type(answer)
```

Output:
```
pandas.core.frame.DataFrame
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3)
```

Output:
```
17974.363
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3)
```

Output:
```
111.559
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python

```

```python

```

```python

```